# Design Decisions with Python Functions



```python
from nose.tools import assert_true, assert_equal, assert_false, assert_almost_equal, assert_raises
```

# Two primary reasons for defining functions:
1. Code reuse: 
    * Write and debug code once. Then I can use this same (correct) code many times
    * This makes upkeep/modifications simpler. When I think of an improved way of implementing something I only need to change it in one location.

2. Procedural Decomposition:
    * A function should do one thing, not multiple things.
    * This can become a matter of style
    

## Function Style According to Mark Thomason
![Mark Thomason](./mark_thomason.jpg)

## The entire function should be visible on your screen within your editor.
## If your function doesn't fit on your screen, get a bigger screen

### What are the implications of these heuristic?
#### Function size changes with age?

## Exercise: Define a function to get a positive integer from a user.
### Requirements
1. Use an infinite while loop
1. Use the input function
1. Keep prompting the user for input until a valid positive integer is provided

We'll use a try/except block to account for users entering non-integer value

```Python
try:
    # Get input from user
    # convert to an integer (this where we could get an exception
    # test for positivity
except ValueError:
    # if we get an input that we can't convert to an integer, we need to do something
```

```python
def get_pos_integer(prompt="Enter a positive integer"):
    while True:
        num = input(prompt)
        try:
            num = int(num)
            if num > 0:
                return num
        except ValueError:
            pass

```

## Does this function do one thing?
## Could we break it into smaller pieces?

### Write a function to test whether a number is positive

<!-- grade_id: cell-b17130446493e1e7, points: 0, grade: false, solution: true -->
```python
def ispositive(x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-163d61c8d4528ae4, points: 0, grade: true, solution: false -->
```python
assert_true(ispositive(5))
assert_false(ispositive(-1))
assert_false(ispositive(0))
```

<!-- grade_id: cell-2514734b675a7247, points: 0, grade: false, solution: true -->
```python
def getint(sint):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c1513c8de2dfd7da, points: 0, grade: true, solution: false -->
```python
assert_equal(getint(543), 543)
assert_raises(ValueError, getint, "4.7")

```

### Using `getint` and `ispositive` rewrite `get_pos_integer`

<!-- grade_id: cell-fa342b9026507d73, points: 0, grade: false, solution: true -->
```python
def get_pos_integer2(prompt="enter a positive integer: "):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e69024ff74b045be, points: 0, grade: true, solution: false -->
```python
# Enter '5'
assert_equal(get_pos_integer2(),5)

# Enter '7'
assert_equal(get_pos_integer2(),7)

```

## Exercise

Following the same style as `get_pos_integer`, write a function `get_value`. `get_value` takes as arguments:

1. A positional argument `converter` that is a function that takes as input a string and returns the desired value
1. A positional argument `tester` that is a function that takes as input a value and returns `True` or `False` depending on whether a desired condition is satisfied.
1. A keyword argument `prompt` that is the prompt to use with `input`.

Test the function with `getint` and `ispositive`.

<!-- grade_id: cell-cebe2a0d20841727, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

```python
get_value(getint, ispositive)
```

## Exercise

Modify `test_ascending` to take in a sequence and test if the elements in `values` are in ascending order.

**Hint:** Use the `all` function.

<!-- grade_id: cell-40b8ef1cd8e54c0b, points: 0, grade: false, solution: true -->
```python
def test_ascending(values):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-519436fcb56d35ad, points: 0, grade: true, solution: false -->
```python
assert_true(test_ascending(("Argos", "Helios", "Zeus")))
assert_false(test_ascending(("Argos", "Zeus", "Helios")))
assert_false(test_ascending(("argos", "Helios", "Zeus")))

```

```python
get_value(get_three_words, 
          test_ascending, 
          prompt="enter three words in ascending alphabetical order separated by spaces: ")

```

## [Avoiding empty list as a default argument](https://stackoverflow.com/questions/1132941/least-astonishment-and-the-mutable-default-argument)

## Observe this unexpected behavior

```python
def foo(a=[]):
    a.append(5)
    return a
print(foo())
print(foo())
```

```python
def i_hate_foo(a=None):
    if a == None:
        a = []
    a.append(5)
    return a
print(i_hate_foo())
print(i_hate_foo())
```

### Variable number of position arguments

A function definition that looks like this

```python

def some_function(*x):
    ### BLOCK OF CODE
```

Has a variable number of positional arguments.

The variable number of positional arguments are passed to the function as a **list.**

```python
def demo1(*x):
    for xx in x:
        print(xx)
demo1(1,"Brian", [1,2,3], {4,5,6})
```

## Exercise

Write a function ``sumthings`` that takes a variable number of arguments and returns the sum of their values.

#### Challenge: Can you do this with a single Python statement within `sumthings`?

```python

```

<!-- grade_id: cell-e7d16b80330ab587, points: 0, grade: false, solution: true -->
```python
def sumthings(*x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-28f6eb94cc2d3bbc, points: 0, grade: true, solution: false -->
```python
assert_equal(sumthings(1,2,3),6)
assert_equal(sumthings(1),1)
assert_equal(sumthings(),0)

```

### Variable number of keyword arguments

A function definition that looks like this

```python

def some_function(**kwargs):
    ### BLOCK OF CODE
```

has a variable number of keyword arguments

The variable number of positional arguments are passed to the function as a **list.**

```python
def demo2(**kwargs):
    for k in kwargs:
        print(k, kwargs[k])
demo2(print="No way", age=29, favorite_number=8, luck="bad")
```