# Practicing Functions

```python
from nose.tools import assert_true, assert_equal, assert_false, assert_almost_equal, assert_raises
```

## Functions to Modify Lists

We stated before that if a function modifies an argument we should always return a `None` value. We can see this in the behavior of the list `reverse` method. `reverse` modifies a list in place, so it should return `None`.

```python
a = [1,3,2]
print(a.reverse())
print(a)
```

## Compare this to the string `upper` method

```python
a = "Tolkien wrote mythology for a modern world."
print(a.upper())
print(a)
```

Output:
```
TOLKIEN WROTE MYTHOLOGY FOR A MODERN WORLD.
Tolkien wrote mythology for a modern world.
```

## The method returns a new string, because `a` is not modified

## Exercise

Write a function `pop` that takes as a positional argument a list (`mylist`) and as a keyword argument index (`index=-1`) and returns two values:

1. The value located at `mylist[index]`
1. A list equal to `mylist` with the value at `index` removed.
1. `mylist` remains **unmodified**

That is, we are implementing a function that recreates the `pop` method of a list but do not modify `mylist`

**Hint:** You could use slicing or you could make a copy of `mylist` inside `pop` (either with slicing or with the [copy](https://docs.python.org/3/library/copy.html) module.

<!-- grade_id: cell-c2718d63cb4d7f12, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d628d4056e1d600a, points: 0, grade: true, solution: false -->
```python
test_list = [1,2,3,4,5]
assert_equal(pop(test_list), (5,[1,2,3,4]))
assert_equal(pop(test_list, index=2), (3,[1, 2, 4, 5]))
assert_equal(pop(test_list, index=0), (1,[ 2, 3, 4, 5]))
assert_equal(pop(test_list, index=-2), (4,[1, 2, 3, 5]))

assert_equal(test_list, [1,2,3,4,5])

```

## Exercise

Write a function that takes as an argument a string and returns a string with each word capitalized.

<!-- grade_id: cell-d45658ce06fb53c7, points: 0, grade: false, solution: true -->
```python
def capitalize(mystring):
    return " ".join([word.capitalize() for word in mystring.split()])

def capitalize(mystring):
    return mystring.title()
```

### How does this work?

The source code for Python is [here](https://github.com/python/cpython). I did a quick search, but couldn't find where `title()` is defined. You can read a brief description in the documentation [here](https://docs.python.org/3/library/stdtypes.html#str.title).

<!-- grade_id: cell-3fccb61666b1dc43, points: 0, grade: true, solution: false -->
```python
assert_equal(capitalize("brian earl chapman"), 'Brian Earl Chapman')
```

```python
help("brian chapman".title)
```

Output:
```
Help on built-in function title:

title(...) method of builtins.str instance
    S.title() -> str
    
    Return a titlecased version of S, i.e. words start with title case
    characters, all remaining cased characters have lower case.
```

#### Since a function is an object, it has attributes

We can find what these attributes are using the `dir` function

```python
dir(len)
```

Output:
```
['__call__',
 '__class__',
 '__delattr__',
 '__dir__',
 '__doc__',
 '__eq__',
 '__format__',
 '__ge__',
 '__getattribute__',
 '__gt__',
 '__hash__',
 '__init__',
 '__init_subclass__',
 '__le__',
 '__lt__',
 '__module__',
 '__name__',
 '__ne__',
 '__new__',
 '__qualname__',
 '__reduce__',
 '__reduce_ex__',
 '__repr__',
 '__self__',
 '__setattr__',
 '__sizeof__',
 '__str__',
 '__subclasshook__',
 '__text_signature__']
```

```python
len.__module__
```

Output:
```
'builtins'
```

```python

```