## Solving the problem using recursion

Recursive functions are functions that call themselves. This is usually done instead of a implementing a loop. Recursion is not usually a great idea in Python, but sometimes it is a good choice and it is good to be aware of.

In designing a recursive function we need to recognize what are our **base cases**: these are the cases where our recursion will stop

### [factorial](https://en.wikipedia.org/wiki/Factorial): A famous recursive function

The factorial function (usually depicted as $n!$), is the product of all non-negative integers less than $n$. $0!$ is defined to be 1 ($0! := 1$).

* The domain of the factorial function is?
* The range of the factorial function is?


#### Before we start, what are our base cases?

```python
def factorial(n):
    pass
```

### [Fibonacci Numbers](https://en.wikipedia.org/wiki/Fibonacci_number): Famous recursive numbers

\begin{eqnarray}
F_n = F_{n-1} + F_{n-2}\\
F_1 = 1\\
F_2 = 1 \\
n \ge 1
\end{eqnarray}



```python
def fibonacci(n):
    pass
```

```python
[fibonacci(n) for n in range(10)]
```

```python
factorial(2)
```

## Modify `get_pos_integer_take2` to get a positive integer from the user using recursion

#### What is our base case?

<!-- grade_id: cell-c44acabc93ef7176, points: 0, grade: false, solution: true -->
```python
def get_pos_integer_take2(prompt="Enter a positive integer"):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
get_pos_integer_take2()
```

```python

```

```python
import random
import seaborn as sns
from nose.tools import assert_equal, assert_true, assert_false, assert_almost_equal
```

Use recursion to write a function `double_every_other` that takes as an argument a sequence of numbers (list or tuple) and returns a tuple with ever other element (starting from the second element) doubled (e.g. [0,1,2,3,4,5]->[0,2,2,6,4,10]).

**Hints:** 

There will be two base cases:

1. When the list is empty
2. When the list has only one element in it

Use slicing and concatenation.

<!-- grade_id: cell-0e669bc20815a9f7, points: 0, grade: false, solution: true -->
```python
def double_every_other(xs):
   pass
```

```python
double_every_other([1,2,3,4,5,6])
```

<!-- grade_id: cell-798764d5e1611e8d, points: 0, grade: true, solution: false -->
```python
assert_equal(double_every_other([1,2,3,4]), [1, 4, 3, 8])
assert_equal(double_every_other([7,3,1,5,0,2,-3,4]), [7, 6, 1, 10, 0, 4, -3, 8])
```

The factorial function (usually denoted with a $!$ as in $5!$) is defined as 
\begin{equation}
n! = 1\times 2 \times \cdots \times (n-1) \times n
\end{equation}

Defining $0!=1$, we can express the factorial function recursively

$$
 n! = 
  \begin{cases} 
   0 & \text{if } n=0 \\
   n\cdot(n-1)!       & \text{otherwise }
  \end{cases}
$$

<!-- grade_id: cell-f1f6960566eac179, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
d
```

<!-- grade_id: cell-3647176a91d5d012, points: 0, grade: true, solution: false -->
```python
assert_equal(factorial(5), 120)
assert_equal(factorial(23),25852016738884976640000)
```

Write a function `product` that computes the product of all the values in a list.

<!-- grade_id: cell-0aa434ef4196b47f, points: 0, grade: false, solution: true -->
```python
def product(xs):
    pass
```

<!-- grade_id: cell-a286bd335ae68929, points: 0, grade: true, solution: false -->
```python
assert_equal(product([1,2,3]),6)
assert_equal(product([1,1,3,4,3,2]),72)
```

Write a function `drop` that takes a list and and an integer N and returns a list with the first N elements dropped.

<!-- grade_id: cell-7aad3abe169ad963, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b4f14c2b0ac62696, points: 0, grade: true, solution: false -->
```python
assert_equal(drop([1,2,3,4,5],2), [3,4,5])
assert_equal(drop([5,4,3,2,1],0), [5, 4, 3, 2, 1])
assert_equal(drop([1,2,3,4,5],9),[])
```

Write a function `zipper` that takes two sequences, `xs` and `ys`, and returns a sequence of the corresponding pairs ---[(x[0],ys[0]), (xs[1],ys[1]),...]---from `xs` and `ys` until one of the sequences is exhausted.

<!-- grade_id: cell-2f8dd69e7c628c8a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e579ddbca10a7fa1, points: 0, grade: true, solution: false -->
```python
assert_equal(zipper([1,2,3],[1,2]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2],[1,2,3]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2,3],[]),[])
assert_equal(zipper([1,2,3,4],[4,3,2,1]), [(1, 4), (2, 3), (3, 2), (4, 1)])
assert_equal(zipper("Brian", "Chapman"), [('B', 'C'), ('r', 'h'), ('i', 'a'), ('a', 'p'), ('n', 'm')])
```

```python

```