# Week 2: Practicing with Functions



```python
from nose.tools import (assert_true, 
                        assert_equal, 
                        assert_false, 
                        assert_almost_equal, 
                        assert_raises)
import math
import pandas as pd
import seaborn as sns
import copy
import matplotlib.pyplot as plt
import numpy as np
import numbers

```

## Class-wide Problem
### Led by Yucheng


A really old way is known as the [Sieve of Eratosthenes](https://en.wikipedia.org/wiki/Sieve_of_Eratosthenes)

Write a function `sieve_prime` that finds primes numbers less than `N` using teh Sieve of Eratosthenes.

The approach here is to create an array of numbers and cross out any number that is a multiple of a prime. As we loop through the array, any number that is not crossed out is a prime.

1. Allocate an array of integers (`nums`) between 2 and N+1
1. Start with i=0
1. Look at the ith number in the array (`nums[i]`), if it is not crossed out it is a prime.
    1. Add `nums[i]` to our list of primes
    1. Cross out all the multiples of `nums[i]`
    1. Increment i

```python
def sieve_prime(N):
    pass
```

```python
assert_equal(sieve_prime(10), [2, 3, 5, 7])
assert_equal(sieve_prime(20), [2, 3, 5, 7, 11, 13, 17, 19])
```

#### Exercise

Write a function `inrange` that takes a positional argument `val` that is a numeric value and a keyword argument `range` that is a two-tuple of lower and upper values defining a range. The function should return `True` if val lies within the range (inclusive) and `False` otherwise.

```python
def inrange(val, range=[0,1]):
    pass
```

```python
assert_true(inrange(0.5))
assert_true(inrange(-math.pi, range=(-6,math.sqrt(2))))
assert_false(inrange(2))
assert_true(inrange(0))
assert_true(inrange(2, (0.5,2)))
```

### Let's write some alternative functions for finding primes

## Exercise:

Write a function `isprime` to check whether a number is prime. 

The function should take an integer as input and return a Boolean.

* Use `if` statements to test for the special cases 
    * $n<2$
    * $n=2$
    * $n\mod2=0$ 
* Use a for loop and the range function 
    * $\text{range}(3,1+\sqrt{n}, 2)$
    * Why $3$? 
    * Why $2$?

```python
def isprime(n):
    pass
```

<!-- grade_id: cell-d7f0ae1402b7514b, points: 0, grade: true, solution: false -->
```python
assert_true(isprime(47))
assert_false(isprime(49))
```

## Exercise

Write a function `get_value`. `get_value` takes as arguments:

1. A positional argument `converter` that is a function that takes as input a string and returns the desired value
1. A positional argument `tester` that is a function that takes as input a value and returns `True` or `False` depending on whether a desired condition is satisfied.
1. A keyword argument `prompt` that is the prompt to use with `input`.

Test the function with `get_number` and `isprime`.

```python
def get_number(x):
    try:
        return int(x)
    except ValueError:
        try:
            return float(x)
        except:
                return complex(x)

```

```python
def get_value(converter, tester, prompt="Enter a prime number"):
    pass
        
    
```

```python
get_value(get_number, isprime)
```

## Exercise

Write a function ``addthings`` that takes a variable number of arguments.

If no arguments are provided, return `None`.

If the first argument is a number, return the sum of all the arguments.

If the first argument is a string, return a concatenation of all the arguments. 

**Hint:** Everything in Python can be converted to a string.

**Note:** This is a problematic function because what is its domain and what is its range?

<!-- grade_id: cell-e7d16b80330ab587, points: 0, grade: false, solution: true -->
```python
def addthings(*xs):
    pass
```

<!-- grade_id: cell-28f6eb94cc2d3bbc, points: 0, grade: true, solution: false -->
```python
assert_equal(addthings("Brian",2,"Wendy"),"Brian2Wendy")
assert_equal(addthings(1),1)
assert_equal(addthings(1,2,3),6)
assert_equal(addthings(),None)

```

### Variable number of keyword arguments

Write a function `get_numeric_items` that takes a variable number of keyword arguments and returns a dictionary of all the arguments with numeric values.

In order to test whether the item is numeric, use the `numbers` library and the `isinstance function`.

```Python
import numbers
isinstance(2.0,numbers.Number)
True

isinstance(5,numbers.Number)
True

isinstance("five",numbers.Number)
False
```

```python
import numbers
isinstance("five",numbers.Number)
```

```python
def get_numeric_items(**kwargs):
    pass
```

```python
assert_equal(get_numeric_items(falling_object="Brian", acceleration=9.8, 
                  acc_units="m/s2",
                  mass=88.4, mass_units="kg",
                  v0=25, v0_units="mph"), 
{'acceleration': 9.8, 'mass': 88.4, 'v0': 25})

assert_equal(get_numeric_items(bowie="david",
                               townsend="pete",
                               young="neil"), 
             {})

assert_equal(get_numeric_items(one=1, two=2, three=3), 
            {'three': 3, 'two': 2, 'one': 1 })
```

## Exercise

Write a function `pop` that takes as a positional argument a list (`mylist`) and as a keyword argument index (`index=-1`) and returns two values:

1. The value located at `mylist[index]`
1. A list equal to `mylist` with the value at `index` removed.
1. `mylist` remains **unmodified**

That is, we are implementing a function that recreates the `pop` method of a list but do not modify `mylist`

**Hint:** You could use slicing or you could make a copy of `mylist` inside `pop` (either with slicing or with the [copy](https://docs.python.org/3/library/copy.html) module.

<!-- grade_id: cell-c2718d63cb4d7f12, points: 0, grade: false, solution: true -->
```python
# Two approaches

def pop(mylist, index=-1):
    pass
```

<!-- grade_id: cell-d628d4056e1d600a, points: 0, grade: true, solution: false -->
```python
test_list = [1,2,3,4,5]
assert_equal(pop(test_list), (5,[1,2,3,4]))
assert_equal(pop(test_list, index=2), (3,[1, 2, 4, 5]))
assert_equal(pop(test_list, index=0), (1,[ 2, 3, 4, 5]))
assert_equal(pop(test_list, index=-2), (4,[1, 2, 3, 5]))

assert_equal(test_list, [1,2,3,4,5])

```

## [Gilbreath's Conjecture](https://en.wikipedia.org/wiki/Gilbreath%27s_conjecture)

>Gilbreath's conjecture is a conjecture in number theory regarding the sequences generated by applying the forward difference operator to consecutive prime numbers and leaving the results unsigned, and then repeating this process on consecutive terms in the resulting sequence, and so forth. The statement is named after mathematician Norman L. Gilbreath who, in 1958, presented it to the mathematical community after observing the pattern by chance while doing arithmetic on a napkin. In 1878, eighty years before Gilbreath's discovery, François Proth had, however, published the same observations along with an attempted proof, which was later shown to be false.

#### Write a function ``seq_diff`` that computes the absolute value of the forward difference of a sequence of primes

**Hints:**
1. The `math` library defines an `fabs` which returns a floating point value for the absolute value.
1. Use the `range` function.
    1. What do you want for the starting value for `range`?
    1. What do you want for the ending value for `range`?
1. How will the length of the output differ from the length of the input?

<!-- grade_id: cell-db3bfe56e4780385, points: 0, grade: false, solution: true -->
```python
def seq_diff(seq):
    pass

```

<!-- grade_id: cell-686cd642a626638c, points: 0, grade: true, solution: false -->
```python
seq = [2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47]
assert_equal(seq_diff(seq), [1, 2, 2, 4, 2, 4, 2, 4, 6, 2, 6, 4, 2, 4])
assert_equal(seq_diff(seq_diff(seq)), [1, 0, 2, 2, 2, 2, 2, 2, 4, 4, 2, 2, 2])
```

```python
seq = [n for n in range(1,50) if isprime(n)]
print(seq)
while True:
    if not seq:
        break   
    seq = seq_diff(seq)
    print(seq)
```

## Playing with Large Measured Data: [Benford's Law](https://en.wikipedia.org/wiki/Benford%27s_law)

### We are going to use census data

```python
pd.read_csv("PEP_2016_PEPANNRSIP.US12A_with_ann.csv", skiprows=1)
```

```python
city_populations = list(pd.read_csv("PEP_2016_PEPANNRSIP.US12A_with_ann.csv", skiprows=1)["Population Estimate (as of July 1) - 2016"])
```

## `city_populations` is a list of integers that are the populations of USA cities

### Create a list of all the populations greater than or equal to 100000

Use list comprehension.

<!-- grade_id: cell-0324eead7bbbaddc, points: 0, grade: false, solution: true -->
```python
city_populations = None 

```

#### Write a function `first_digits`

`first_digits` takes a list of numbers and returns a list of integers corresponding to the first digit (left most, most significant) for each of those numbers.

```python
def first_digits(xs):
    pass
```

### Create a list named `city_populations_first_digits` consisting of the first digit in each number converted to an int

<!-- grade_id: cell-007a102e12198872, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION
city_populations_first_digits = first_digits(city_populations)
### END SOLUTION
```

```python
from collections import Counter
```

### What does the distribution of first digits look like?

#### It should look like this
![digit distribution](./population_digits.png)

```python
sns.distplot(city_populations_first_digits, kde=False)
plt.savefig("population_digits.png")
```

## How does this differ from randomly generated numbers?

### We have to duplicate all our prior code...
### ...Unless we put all the functionality in functions.

The resulting distribution should look something this:

![random digits](./random_digits.png)

```python
random_numbers = np.random.random_integers(99999, 10000000, len(city_populations_strings))
```

```python
sns.distplot(first_digits(random_numbers), kde=False)

```

```python

```

#### Write a lambda expression

Write a lambda expression for the function

\begin{equation}
f(x) = \frac{x}{1+x^2}
\end{equation}

Use this lambda expression, list comprehension, and the numpy [arange](https://docs.scipy.org/doc/numpy/reference/generated/numpy.arange.html) function to create a list of values for this function for $x$ values between 1 (inclusive) and 2 (exclusive) in 0.05 increments.my

```python
my_result = None

```

```python
assert_true(np.allclose(my_result,
                        [0.5, 0.49940546967895366, 0.49773755656108604, 0.49515608180839615, 
                         0.49180327868852458, 0.48780487804878042, 0.48327137546468402, 
                         0.47829937998228522, 0.47297297297297297, 0.46736502820306197, 
                         0.46153846153846151, 0.45554739162380598, 0.44943820224719094, 
                         0.44325050369375413, 0.43701799485861176, 0.43076923076923074, 
                         0.42452830188679241, 0.41831543244771047, 0.41214750542299339, 
                         0.40603852160333148]))
```

### Write a lambda expression

Use a lambda expression to sort the list data by the sum of the values for each tuple

```python
data = [(1,(1,2,3)),(2,(3,2)),(3,(4,5,-1,5,6)),(4,{1,2,3})]
```

```python
data.sort(key=None)
```

### [Fibonacci Numbers](https://en.wikipedia.org/wiki/Fibonacci_number): Famous recursive numbers

\begin{eqnarray}
F_n = F_{n-1} + F_{n-2}\\
F_1 = 1\\
F_2 = 1 \\
n \ge 1
\end{eqnarray}



```python
def fibonacci(n):
    pass
```

```python
assert_equal([fibonacci(n) for n in range(1,10)],
             [1, 1, 2, 3, 5, 8, 13, 21, 34])
```

Write a function `drop` that takes a sequence and and an integer N and returns the sequence with the first N elements dropped.

<!-- grade_id: cell-7aad3abe169ad963, points: 0, grade: false, solution: true -->
```python
def drop(xs,n):
    pass
        
```

<!-- grade_id: cell-b4f14c2b0ac62696, points: 0, grade: true, solution: false -->
```python
assert_equal(drop([1,2,3,4,5],2), [3,4,5])
assert_equal(drop([5,4,3,2,1],0), [5, 4, 3, 2, 1])
assert_equal(drop([1,2,3,4,5],9),[])
```

Write a function `zipper` that takes two sequences, `xs` and `ys`, and returns a sequence of the corresponding pairs ---[(x[0],ys[0]), (xs[1],ys[1]),...]---from `xs` and `ys` until one of the sequences is exhausted.

<!-- grade_id: cell-2f8dd69e7c628c8a, points: 0, grade: false, solution: true -->
```python
def zipper(xs, ys):
    pass
```

<!-- grade_id: cell-e579ddbca10a7fa1, points: 0, grade: true, solution: false -->
```python
assert_equal(zipper([1,2,3],[1,2]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2],[1,2,3]), [(1, 1), (2, 2)])
assert_equal(zipper([1,2,3],[]),[])
assert_equal(zipper([1,2,3,4],[4,3,2,1]), [(1, 4), (2, 3), (3, 2), (4, 1)])
assert_equal(zipper("Brian", "Chapman"), [('B', 'C'), ('r', 'h'), ('i', 'a'), ('a', 'p'), ('n', 'm')])
```

```python

```

```python

```